### Vowel Count [ungraded]
Write a Python function `vowelsCount(s)` that returns the total number of uppercase and lowercase vowels in string `s`. 

<!-- grade_id: cell-f39e8634b7bc8e7f, points: 0, grade: false, solution: true -->
```python
def vowelsCount(s):
    cnt = 0
    for c in s:
        if c in "aeiouAEIOU":
            cnt += 1
    return cnt
```

You can use the cells below for testing and you can add cells with your own test cases.

<!-- grade_id: cell-72293d74ab104d94, points: 0, grade: true, solution: false -->
```python
assert vowelsCount("January") == 3
```

<!-- grade_id: cell-634cf96cb3492601, points: 0, grade: true, solution: false -->
```python
assert vowelsCount("April") == 2
```

<!-- grade_id: cell-d1051f8061fda718, points: 0, grade: true, solution: false -->
```python
assert vowelsCount("abcdefghijklmnopqrtsuvwxyz") == 5
```

<!-- grade_id: cell-0cfbaf282c1ce854, points: 0, grade: true, solution: false -->
```python
assert vowelsCount("ABCDEFGHIJKLMNOPQRTSUVWXYZ") == 5
```

<!-- grade_id: cell-c53f52b4fd864f45, points: 0, grade: true, solution: false -->
```python
assert vowelsCount("") == 0
```

```python

```