### Amicable Numbers [4 points]

Let `d(n)` be the sum of proper divisors of `n` (numbers less than `n` which divide evenly into `n`). For example, the proper divisors of `6` are `1`, `2`, `3`, therefore `d(6) = 1 + 2 + 3 = 6`.  The proper divisors of `14` are `1`, `2`, `7`, therefore `d(14) = 1 + 2 + 7 = 10`. If `d(a) = b` and `d(b) = a`, where `a ≠ b`, then `a` and `b` are an amicable pair and each of `a` and `b` are called *amicable numbers*. For example, the proper divisors of `220` are `1`, `2`, `4`, `5`, `10`, `11`, `20`, `22`, `44`, `55` and `110`, therefore `d(220) = 284`. The proper divisors of `284` are `1`, `2`, `4`, `71` and `142`, therefore `d(284) = 220`. Thus, `220` and `284` are amicable numbers.

Let `a(m)` be the sum of all amicable numbers less than `m` for `m ≥ 1`. Implement the functions `d(n)` and `a(m)` in Python by completing the bodies below. Run the cells with tests to check your implementation. You can add cells with your own tests. The grade is being automatically determined based on the provided tests and further hidden tests. The efficiency is not being judged, however, if your implementation takes longer than 30 seconds, the tests will time out and fail.

<!-- grade_id: cell-978aa12127859e46, points: 0, grade: false, solution: true -->
```python
def d(n):  
    """sum of divisors of n for n ≥ 1"""
    total = 0
    for i in range(1,n):
        if n%i == 0:
            total += i

    return total
```

<!-- grade_id: cell-71d527cfeb988681, points: 2, grade: true, solution: false -->
```python
assert d(1) == 0
assert d(220) == 284
assert d(284) == 220
assert d(999) == 521
assert d(1184) == 1210
assert d(1210) == 1184
```

<!-- grade_id: cell-fe644888e4ee1e0f, points: 0, grade: false, solution: true -->
```python
def a(m): 
    """sum of amicable numbers less than m for m ≥ 1"""
    amicable_numbers = set()
    for i in range(1,m):
        num = d(i)
        if d(num) == i and num != i:
            amicable_numbers.add(num)
    return sum(amicable_numbers)

```

<!-- grade_id: cell-baf0fee5afc65b8b, points: 2, grade: true, solution: false -->
```python
assert a(1000) == 504
assert a(5000) == 8442
```

_Hint:_ Each of the two Python functions can be implemented in a single line using [generator expressions](https://docs.python.org/3/glossary.html#term-generator-expression).

```python
assert a(0) == 0
```