### Number of Interleavings [4 points]

Suppose process `A` executes in a  single step (atomic instruction) and `B` executes in two steps. Then the execution of `A ‖ B` leads to three possible interleavings:

| Process | Steps         |
| :------ | :------------ |
| `A`     | ➀             |
| `B`     | ➊ ➋           |
| `A ‖ B` | ➀ ➊ ➋ <br> ➊ ➀ ➋ <br> ➊ ➋ ➀ |

Suppose processes `A` and `B` execute each of a two steps. Then the execution of `A ‖ B` leads to six possible interleavings:

| Process | Steps         |
| :------ | :------------ |
| `A`     | ➀ ➁          |
| `B`     | ➊ ➋          |
| `A ‖ B` | ➀ ➁ ➊ ➋ <br> ➀ ➊ ➁ ➋ <br> ➀ ➊ ➋ ➁ <br> ➊ ➀ ➁ ➋ <br> ➊ ➀ ➋ ➁ <br> ➊ ➋ ➀ ➁ |

The table is constructed by first taking ➀ as the first step and enumerating all interleavings of ➁ and ➊ ➋, then taking ➊ as the first step and enumerating all interleaving of ➀ ➁ and ➋.

Now consider the example from the course notes with processes `inc1` and `inc2`. Each of these execute in three steps, abbreviated ➀ ➁ ➂ and ➊ ➋ ➌. How many possible interleavings of `inc1 ‖ inc2` are there? Enumerate them!

<!-- grade_id: cell-edc8d3f3c30742b6, points: 0, grade: true, solution: true -->
1.  ➀ ➁ ➂ ➊ ➋ ➌
2.  ➀ ➁ ➊ ➂ ➋ ➌
3.  ➀ ➁ ➊ ➋ ➂ ➌
4.  ➀ ➁ ➊ ➋ ➌ ➂
5.  ➀ ➊ ➁ ➂ ➋ ➌
6.  ➀ ➊ ➁ ➋ ➂ ➌
7.  ➀ ➊ ➁ ➋ ➌ ➂
8.  ➀ ➊ ➋ ➁ ➂ ➌
9.  ➀ ➊ ➋ ➁ ➌ ➂
10. ➀ ➊ ➋ ➌ ➁ ➂
11. ➊ ➋ ➌ ➀ ➁ ➂
12. ➊ ➋ ➀ ➌ ➁ ➂
13. ➊ ➋ ➀ ➁ ➌ ➂
14. ➊ ➋ ➀ ➁ ➂ ➌
15. ➊ ➀ ➋ ➌ ➁ ➂
16. ➊ ➀ ➋ ➁ ➌ ➂
17. ➊ ➀ ➋ ➁ ➂ ➌
18. ➊ ➀ ➁ ➋ ➌ ➂
19. ➊ ➀ ➁ ➋ ➂ ➌
20. ➊ ➀ ➁ ➂ ➋ ➌

**Bonus Question [4 points].** Write a Python function `interleavings(a, b)` that takes two strings (or lists), `a` and `b`, and returns a set with all possible interleavings of the elements of `a` and `b`. The grade is being automatically determined by the provided and hidden tests.

<!-- grade_id: cell-da52e784998e0527, points: 0, grade: true, solution: true -->
```python
def interleavings(a, b):
    # len(a) > 0 and len(b) > 0
    if a == "":
        return {b}
    if b == "":
        return {a}
        
    results = set()
    for resultA in interleavings(a[1:], b):
        results.add(a[0] + resultA)
    for resultB in interleavings(a, b[1:]):
        results.add(b[0] + resultB)

    return results
```

<!-- grade_id: cell-7e00e808749e784a, points: 4, grade: true, solution: false -->
```python
assert interleavings('➀', '➊') == {'➀➊', '➊➀'}
assert interleavings('➀', '➊➋') == {'➀➊➋', '➊➀➋', '➊➋➀'}
assert interleavings('➀➁', '➊➋') == {'➀➁➊➋', '➀➊➁➋', '➀➊➋➁', '➊➀➁➋', '➊➀➋➁', '➊➋➀➁'}
assert interleavings('➀', '➊➋➌') == {'➀➊➋➌', '➊➀➋➌', '➊➋➀➌', '➊➋➌➀'}
```