### Evaluating a Polynomial [7 points]

A polynomial in one variable (*univariate*) is of the form `p(x) = c(n) × xⁿ + c(n-1) × xⁿ⁻¹ + ... + c(1) × x + c(0)`, where the `c(i)` are the _coefficients_ and `n` is the _degree_.

Assuming that `N ≥ 0`, the array `c: 0 .. N – 1 → integer` has the coefficients, and `x` is an integer, write a program that, for integer variable `s`, establishes the postcondition

```algorithm
s = (∑ i ∈ 0 .. N - 1 • a(i) × xⁱ)
```

##### 1. With exponentiation [4 points]

The program should contain only one loop and use exponentiation. Give the programs textually and add the appropriate annotation! Use as a template:

```algorithm
{true}
s = 0
k = 0
while k < N do:
    s, k = s + a(k) × x^k, k + 1
{s = (∑ i ∈ 0 .. N – 1 • a(i) × xⁱ)}
```

Give the correctness conditions and prove them!

<!-- grade_id: cell-0ad0faf785be6234, points: 4, grade: true, solution: true -->
<span style = "float:left;border-right:1em solid transparent">(1)</span>

```algorithm
{true}
s, k := 0, 0
{s = (∑ i ∈ 0 .. k – 1 • a(i)*x^i) ∧ 0 ≤ k ≤ N}
```
proof:
```algorithm
    (s = (∑ i ∈ 0 .. k – 1 • a(i)*x^i) ∧ 0 ≤ k ≤ N)[s, k := 0, 0]
≡        {by substitution, arithmetic}
    0 = (∑ i ∈ 0 .. – 1 • a(i)*x^i) ∧ 0 ≤ 0 ≤ N
≡        {empty sum is 0, arithmetic}
    0 = 0 ∧ 0 ≤ N
≡        {assumption N ≥ 0}
    true
```

<span style = "float:left;border-right:1em solid transparent">(2)</span>

```algorithm
{s = (∑ i ∈ 0 .. k – 1 • a(i)*x^i) ∧ 0 ≤ k ≤ N}
while k < N do
    s, k := s + a(k)*x^k, k + 1
{s = (∑ i ∈ 0 .. N – 1 • a(i)*x^i)}
```

<span style = "float:left;border-right:1em solid transparent">(2.1)</span>

```algorithm
s = (∑ i ∈ 0 .. k – 1 • a(i)*x^i) ∧ 0 ≤ k ≤ N ∧ k < N ⇒ (s = (∑ i ∈ 0 .. k – 1 • a(i)*x^i) ∧ 0 ≤ k ≤ N)[s, k := s + a(k)*x^k, k + 1]

```
proof:

```algorithm
    (s = (∑ i ∈ 0 .. k – 1 • a(i)*x^i) ∧ 0 ≤ k ≤ N)[s, k := s + a(k)*x^k, k + 1]
≡        {by substitution}
    s + a(k)*x^k = (∑ i ∈ 0 .. (k + 1) – 1 • a(i)*x^i) ∧ 0 ≤ k + 1 ≤ N
⇐        {by arithmetic}
    s + a(k)*x^k = (∑ i ∈ 0 .. k • a(i)*x^i) ∧ 0 ≤ k < N
≡        {by domain split with i < k and i = k}
    s + a(k)*x^k = (∑ i ∈ 0 .. k – 1 • a(i)*x^i) + a(k)*x^k ∧ 0 ≤ k < N
≡        {by arithmetic}
    s = (∑ i ∈ 0 .. k – 1 • a(i)*x^i) ∧ 0 ≤ k < N
≡        {by arithmetic}
    s = (∑ i ∈ 0 .. k – 1 • a(i)*x^i) ∧ 0 ≤ k ≤ N ∧ k < N
```

<span style = "float:left;border-right:1em solid transparent">(2.2)</span>

```algorithm
s = (∑ i ∈ 0 .. k – 1 • a(i)*x^i) ∧ 0 ≤ k ≤ N ∧ k ≥ N ⇒ s = (∑ i ∈ 0 .. N – 1 • a(i)*x^i)
```

proof:

```algorithm
    s = (∑ i ∈ 0 .. k – 1 • a(i)*x^i) ∧ 0 ≤ k ≤ N ∧ k ≥ N ⇒
⇒        {by arithmetic}
    s = (∑ i ∈ 0 .. k – 1 • a(i)*x^i) ∧ k = N
≡        {by substitution of equals for equals}
    s = (∑ i ∈ 0 .. N – 1 • a(i)*x^i) ∧ k = N
⇒        {logic}
    s = (∑ i ∈ 0 .. N – 1 • a(i)*x^i)
```
)
```

```

##### 2. Without exponentiation [2 points]

The program should contain only one loop and use only addition and multiplication, no exponentiation. Add appropriate annotation. You do not need to prove that the program is correct, but the loop invariant has to be strong enough to allow a proof (you may want to carry out the proofs for yourself). Use the above template.

<!-- grade_id: cell-afb351faef05f180, points: 2, grade: true, solution: true -->
```algorithm
{true}
var power = 1              // x^0 = 1
for i = 0 to N - 1 do:
    s += a(i) × xPower
    power = power × x
{s = (∑ i ∈ 0 .. N – 1 • a(i) × xⁱ)}
```

##### 3. Implementation [1 point]

Implement your algorithm without exponentions in Python!

<!-- grade_id: cell-293ee825ca9be7d7, points: 0, grade: true, solution: true -->
```python
def poly(a, x):
    s = 0
    power = 1

    for i in a:
        s += i * power
        power *= x

    return s;
```

<!-- grade_id: cell-a563d7766ede7e8a, points: 1, grade: true, solution: false -->
```python
assert poly([1, 2, 3], 3) == 34
assert poly([7], 6) == 7
assert poly([], 3) == 0
assert poly([2, 0, 0], 7) == 2
assert poly([0, 0, 2], 7) == 98
```

```python

```