### Tree Traversal with Generators (Python)

Suppose we are given a tree structure and want to process each leaf one by one. The tree is represented by nested tuples (or lists) in Python, like `((1, 2), 3, (4, (5, 6)))`. The natural way is to traverse a tree recursively to visit the leaves in order, as in:

```python
from collections.abc import Iterable
```

```python
def printleaves(node):
    if isinstance(node, Iterable): # internal node
        for c in node: printleaves(c)
    else: print(node) # leaf
```

```python
printleaves(((1, 2), 3, (4, (5, 6))))
```

Output:
```
1
2
3
4
5
6
```

Type `Iterable` is an *abstract base class* that allows objects to be used in `for` loops. The type test `isinstance(node, Iterable)` checks if `node` is composed and can be iterated over, i.e. is an internal node or is atomic:

```python
isinstance(5, Iterable), isinstance([1, 2], Iterable), isinstance((3, 4, 5), Iterable)
```

Output:
```
(False, True, True)
```

The task is to define the Python function `leaves` as a generator that will return the leaves in sequence from left to right without explicitly constructing a sequence with the leaves:

<!-- grade_id: cell-ff2e3d55a18e9cdb, points: 0, grade: true, solution: true -->
```python
def leaves(tree):
    if isinstance(tree, Iterable):
        for c in tree: yield from leaves(c)
    else: yield tree
```

*Hint:* The solution has about three lines of code; if you use [`yield from`](https://docs.python.org/3/whatsnew/3.3.html#pep-380), it can be further shortened. Use the cells below for testing.

```python
l = leaves(((1, 2), 3, (4, (5, 6))))
next(l), next(l), next(l), next(l)
```

Output:
```
(1, 2, 3, 4)
```

Above cell should output `(1, 2, 3, 4)`. Your implementation must pass the test below: [4 points]

<!-- grade_id: cell-4fa5b8b3a3c73576, points: 4, grade: true, solution: false -->
```python
assert list(leaves(((1, 2), 3, (4, (5, 6))))) == [1, 2, 3, 4, 5, 6]
assert sum(leaves(((1, 2), 3, (4, (5, 6))))) == 21
```

```python

```